rule_files:
  - /tmp/prometheus-rules-test.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: rox_central_postgres_connected{namespace="rhacs-1234", pod="central-1234", container="central"}
        values: "1+0x10 0+0x50"
    alert_rule_test:
      - eval_time: 15m
        alertname: RHACSCentralPostgresConnectionDown
        exp_alerts: []
      - eval_time: 40m
        alertname: RHACSCentralPostgresConnectionDown
        exp_alerts:
          - exp_labels:
              alertname: RHACSCentralPostgresConnectionDown
              pod: central-1234
              container: central
              namespace: rhacs-1234
              severity: critical
            exp_annotations:
              summary: "Central container `central-1234/central` database connection in namespace `rhacs-1234` is down or is in a bad shape."
              description: "Central container `central-1234/central` database connection in namespace `rhacs-1234` has been down or is in a bad shape for at least 10 minutes."
              sop_url: "https://gitlab.cee.redhat.com/stackrox/acs-managed-service-runbooks/blob/master/sops/dp-010-rhacs-instance-db-unavailable.md"
