rule_files:
  - /tmp/prometheus-rules-test.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: up{namespace="rhacs-1234", pod="scanner-1234-5678", instance="1.2.3.4:9090"}
        values: "0+0x20 1+0x20"
      - series: kube_pod_container_status_ready{namespace="rhacs-1234", pod="scanner-1234-5678", container="scanner"}
        values: "1+0x40"
    alert_rule_test:
      - eval_time: 10m
        alertname: RHACSScannerScrapeFailed
        exp_alerts: []
      - eval_time: 25m
        alertname: RHACSScannerScrapeFailed
        exp_alerts:
          - exp_labels:
              alertname: RHACSScannerScrapeFailed
              instance: 1.2.3.4:9090
              namespace: rhacs-1234
              pod: scanner-1234-5678
              severity: critical
            exp_annotations:
              summary: "Prometheus unable to scrape metrics from target `scanner-1234-5678` in namespace `rhacs-1234`."
              description: "During the last 10 minutes, only `45.45%` of scrapes of target `scanner-1234-5678` in namespace `rhacs-1234` were successful. This alert is raised when less than 50% of scrapes are successful."
              sop_url: "https://github.com/stackrox/acs-managed-service-runbooks/blob/master/sops/dp-003-rhacs-instance-unavailable.md"
