rule_files:
  - /tmp/prometheus-rules-test.yaml

evaluation_interval: 30s

group_eval_order:
  - rhacs-central.sli
  - rhacs-central.slo
  - rhacs-central.alerts

tests:
  # Ensure that a service with zero requests doesn't fire an alert
  - interval: 5m
    input_series:
      - series: kube_deployment_status_replicas_ready{deployment="central", namespace="rhacs-aaaabbbbccccddddeeee"}
        values: "1x1000"
      - series: grpc_server_handled_total{job="central", grpc_type="unary", grpc_code="OK", namespace="rhacs-aaaabbbbccccddddeeee", rhacs_instance_id="aaaabbbbccccddddeeee"}
        values: "1x1000"
      - series: haproxy_backend_http_responses_total{job="router_internal_default", code="2xx", exported_namespace="rhacs-aaaabbbbccccddddeeee"}
        values: "1x1000"
      - series: grpc_server_handled_total{job="central", grpc_type="unary", grpc_code="Unavailable", namespace="rhacs-aaaabbbbccccddddeeee", rhacs_instance_id="aaaabbbbccccddddeeee"}
        values: "0x1000"
      - series: haproxy_backend_http_responses_total{job="router_internal_default", code="5xx", exported_namespace="rhacs-aaaabbbbccccddddeeee"}
        values: "0x1000"
    alert_rule_test:
      - eval_time: 500m
        alertname: Central availability error budget exhaustion - 50%
        exp_alerts: []
  # Ensure that a service with zero health checks doesn't fire an alert
  - interval: 5m
    input_series:
      - series: kube_deployment_status_replicas_ready{deployment="central", namespace="rhacs-aaaabbbbccccddddeeee"}
        values: "1x400"
      - series: grpc_server_handled_total{job="central", grpc_type="unary", grpc_code="OK", namespace="rhacs-aaaabbbbccccddddeeee", rhacs_instance_id="aaaabbbbccccddddeeee"}
        values: "4+4x400"
      - series: haproxy_backend_http_responses_total{job="router_internal_default", code="2xx", exported_namespace="rhacs-aaaabbbbccccddddeeee"}
        values: "1x1000"
      - series: grpc_server_handled_total{job="central", grpc_type="unary", grpc_code="Unavailable", namespace="rhacs-aaaabbbbccccddddeeee", rhacs_instance_id="aaaabbbbccccddddeeee"}
        values: "0x1000"
      - series: haproxy_backend_http_responses_total{job="router_internal_default", code="5xx", exported_namespace="rhacs-aaaabbbbccccddddeeee"}
        values: "0x1000"
    alert_rule_test:
      - eval_time: 500m
        alertname: Central availability error budget exhaustion - 50%
        exp_alerts: []
      - eval_time: 500m
        alertname: Central availability error budget exhaustion - 90%
        exp_alerts: [ ]
  # Ensure that a service with zero requests but the pod is down *does* fire an alert
  - interval: 5m
    input_series:
      - series: kube_deployment_status_replicas_ready{deployment="central", namespace="rhacs-aaaabbbbccccddddeeee"}
        values: "0x1000"
      - series: grpc_server_handled_total{job="central", grpc_type="unary", grpc_code="OK", namespace="rhacs-aaaabbbbccccddddeeee", rhacs_instance_id="aaaabbbbccccddddeeee"}
        values: "1x1000"
      - series: haproxy_backend_http_responses_total{job="router_internal_default", code="2xx", exported_namespace="rhacs-aaaabbbbccccddddeeee"}
        values: "1x1000"
      - series: grpc_server_handled_total{job="central", grpc_type="unary", grpc_code="Unavailable", namespace="rhacs-aaaabbbbccccddddeeee", rhacs_instance_id="aaaabbbbccccddddeeee"}
        values: "0x1000"
      - series: haproxy_backend_http_responses_total{job="router_internal_default", code="5xx", exported_namespace="rhacs-aaaabbbbccccddddeeee"}
        values: "0x1000"
      - series: probe_success{app="rhacs", endpoint="central-ui", namespace="rhacs-aaaabbbbccccddddeeee", rhacs_instance_id="aaaabbbbccccddddeeee"}
        values: "1x1000"
      - series: probe_success{app="rhacs", endpoint="central-data", namespace="rhacs-aaaabbbbccccddddeeee", rhacs_instance_id="aaaabbbbccccddddeeee"}
        values: "1x1000"
    alert_rule_test:
      - eval_time: 500m
        alertname: Central availability error budget exhaustion - 90%
        exp_alerts:
          - exp_labels:
              alertname: Central availability error budget exhaustion - 90%
              service: central
              severity: critical
              namespace: rhacs-aaaabbbbccccddddeeee
              rhacs_instance_id: aaaabbbbccccddddeeee
            exp_annotations:
              message: "High availability error budget exhaustion for central. Current exhaustion: 124.1%."
              sop_url: "https://gitlab.cee.redhat.com/stackrox/acs-managed-service-runbooks/blob/master/sops/dp-018-rhacs-central-slo-alerts.md"
  # Ensure that a service with zero requests but health checks fail *does* fire an alert
  - interval: 5m
    input_series:
      - series: kube_deployment_status_replicas_ready{deployment="central", namespace="rhacs-6ccace6f8445fc1c8306"}
        values: "1x1000"
      - series: grpc_server_handled_total{job="central", grpc_type="unary", grpc_code="OK", namespace="rhacs-6ccace6f8445fc1c8306", rhacs_instance_id="6ccace6f8445fc1c8306"}
        values: "1x1000"
      - series: haproxy_backend_http_responses_total{job="router_internal_default", code="2xx", exported_namespace="rhacs-6ccace6f8445fc1c8306"}
        values: "1x1000"
      - series: grpc_server_handled_total{job="central", grpc_type="unary", grpc_code="Unavailable", namespace="rhacs-6ccace6f8445fc1c8306", rhacs_instance_id="6ccace6f8445fc1c8306"}
        values: "0x1000"
      - series: haproxy_backend_http_responses_total{job="router_internal_default", code="5xx", exported_namespace="rhacs-6ccace6f8445fc1c8306"}
        values: "0x1000"
      - series: probe_success{app="rhacs", endpoint="central-ui", namespace="rhacs-6ccace6f8445fc1c8306", rhacs_instance_id="6ccace6f8445fc1c8306"}
        values: "0x500 1x500"
      - series: probe_success{app="rhacs", endpoint="central-data", namespace="rhacs-6ccace6f8445fc1c8306", rhacs_instance_id="6ccace6f8445fc1c8306"}
        values: "1x500 0x500"
    alert_rule_test:
      - eval_time: 500m
        alertname: Central availability error budget exhaustion - 90%
        exp_alerts:
          - exp_labels:
              alertname: Central availability error budget exhaustion - 90%
              service: central
              severity: critical
              namespace: rhacs-6ccace6f8445fc1c8306
              rhacs_instance_id: 6ccace6f8445fc1c8306
            exp_annotations:
              message: "High availability error budget exhaustion for central. Current exhaustion: 124.1%."
              sop_url: "https://gitlab.cee.redhat.com/stackrox/acs-managed-service-runbooks/blob/master/sops/dp-018-rhacs-central-slo-alerts.md"
  # Central availability error budget exhaustion - 90%
  - interval: 5m
    input_series:
      # 200m downtime due to pod not ready. Out of 28 days, this equates to ~0.5% downtime.
      - series: kube_deployment_status_replicas_ready{deployment="central", namespace="rhacs-aaaabbbbccccddddeeee"}
        values: "1+0x260 0+0x40 1+0x100"
      - series: grpc_server_handled_total{job="central", grpc_type="unary", grpc_code="OK", namespace="rhacs-aaaabbbbccccddddeeee", rhacs_instance_id="aaaabbbbccccddddeeee"}
        values: "1+1x400"
      - series: haproxy_backend_http_responses_total{job="router_internal_default", code="2xx", exported_namespace="rhacs-aaaabbbbccccddddeeee"}
        values: "4+4x400"
      # 200m downtime due Unavailable or 5xx responses. Out of 28 days, this equates to ~0.5% downtime.
      - series: grpc_server_handled_total{job="central", grpc_type="unary", grpc_code="Unavailable", namespace="rhacs-aaaabbbbccccddddeeee", rhacs_instance_id="aaaabbbbccccddddeeee"}
        values: "0+0x360 0+1x40"
      - series: haproxy_backend_http_responses_total{job="router_internal_default", code="5xx", exported_namespace="rhacs-aaaabbbbccccddddeeee"}
        values: "0+0x360 0+4x40"
      - series: probe_success{app="rhacs", endpoint="central-ui", namespace="rhacs-aaaabbbbccccddddeeee", rhacs_instance_id="aaaabbbbccccddddeeee"}
        values: "1x400"
      - series: probe_success{app="rhacs", endpoint="central-data", namespace="rhacs-aaaabbbbccccddddeeee", rhacs_instance_id="aaaabbbbccccddddeeee"}
        values: "1x400"
    alert_rule_test:
      - eval_time: 100m
        alertname: Central availability error budget exhaustion - 90%
        exp_alerts: []
      - eval_time: 2000m
        alertname: Central availability error budget exhaustion - 90%
        exp_alerts:
          - exp_labels:
              alertname: Central availability error budget exhaustion - 90%
              service: central
              severity: critical
              namespace: rhacs-aaaabbbbccccddddeeee
              rhacs_instance_id: aaaabbbbccccddddeeee
            exp_annotations:
              message: "High availability error budget exhaustion for central. Current exhaustion: 97.97%."
              sop_url: "https://gitlab.cee.redhat.com/stackrox/acs-managed-service-runbooks/blob/master/sops/dp-018-rhacs-central-slo-alerts.md"

  # Central availability error budget exhaustion - 70%
  - interval: 5m
    input_series:
      # 175m downtime due to pod not ready. Out of 28 days, this equates to ~0.43% downtime.
      - series: kube_deployment_status_replicas_ready{deployment="central", namespace="rhacs-ffffgggghhhhiiiijjjj"}
        values: "1+0x265 0+0x35 1+0x100"
      - series: grpc_server_handled_total{job="central", grpc_type="unary", grpc_code="OK", namespace="rhacs-ffffgggghhhhiiiijjjj", rhacs_instance_id="ffffgggghhhhiiiijjjj"}
        values: "1+1x400"
      - series: haproxy_backend_http_responses_total{job="router_internal_default", code="2xx", exported_namespace="rhacs-ffffgggghhhhiiiijjjj"}
        values: "4+4x400"
      # 175m downtime due Unavailable or 5xx responses. Out of 28 days, this equates to ~0.43% downtime.
      - series: grpc_server_handled_total{job="central", grpc_type="unary", grpc_code="Unavailable", namespace="rhacs-ffffgggghhhhiiiijjjj", rhacs_instance_id="ffffgggghhhhiiiijjjj"}
        values: "0+0x365 0+1x35"
      - series: haproxy_backend_http_responses_total{job="router_internal_default", code="5xx", exported_namespace="rhacs-ffffgggghhhhiiiijjjj"}
        values: "0+0x365 0+4x35"
      - series: probe_success{app="rhacs", endpoint="central-ui", namespace="rhacs-aaaabbbbccccddddeeee", rhacs_instance_id="aaaabbbbccccddddeeee"}
        values: "1x400"
      - series: probe_success{app="rhacs", endpoint="central-data", namespace="rhacs-aaaabbbbccccddddeeee", rhacs_instance_id="aaaabbbbccccddddeeee"}
        values: "1x400"
      - series: probe_success{app="rhacs", endpoint="central-ui", namespace="rhacs-ffffgggghhhhiiiijjjj", rhacs_instance_id="ffffgggghhhhiiiijjjj"}
        values: "1x400"
      - series: probe_success{app="rhacs", endpoint="central-data", namespace="rhacs-ffffgggghhhhiiiijjjj", rhacs_instance_id="ffffgggghhhhiiiijjjj"}
        values: "1x400"
    alert_rule_test:
      - eval_time: 100m
        alertname: Central availability error budget exhaustion - 70%
        exp_alerts: []
      - eval_time: 2000m
        alertname: Central availability error budget exhaustion - 70%
        exp_alerts:
          - exp_labels:
              alertname: Central availability error budget exhaustion - 70%
              service: central
              severity: warning
              namespace: rhacs-ffffgggghhhhiiiijjjj
              rhacs_instance_id: ffffgggghhhhiiiijjjj
            exp_annotations:
              message: "High availability error budget exhaustion for central. Current exhaustion: 85.57%."
              sop_url: "https://gitlab.cee.redhat.com/stackrox/acs-managed-service-runbooks/blob/master/sops/dp-018-rhacs-central-slo-alerts.md"

  # Central availability error budget exhaustion - 50%
  - interval: 5m
    input_series:
      # 105m downtime due to pod not ready. Out of 28 days, this equates to ~0.25% downtime.
      - series: kube_deployment_status_replicas_ready{deployment="central", namespace="rhacs-kkkkllllmmmmnnnnoooo"}
        values: "1+0x279 0+0x21 1+0x100"
      - series: grpc_server_handled_total{job="central", grpc_type="unary", grpc_code="OK", namespace="rhacs-kkkkllllmmmmnnnnoooo", rhacs_instance_id="kkkkllllmmmmnnnnoooo"}
        values: "1+1x400"
      - series: haproxy_backend_http_responses_total{job="router_internal_default", code="2xx", exported_namespace="rhacs-kkkkllllmmmmnnnnoooo"}
        values: "4+4x400"
      # 105m downtime due Unavailable or 5xx responses. Out of 28 days, this equates to ~0.25% downtime.
      - series: grpc_server_handled_total{job="central", grpc_type="unary", grpc_code="Unavailable", namespace="rhacs-kkkkllllmmmmnnnnoooo", rhacs_instance_id="kkkkllllmmmmnnnnoooo"}
        values: "0+0x379 0+1x21"
      - series: haproxy_backend_http_responses_total{job="router_internal_default", code="5xx", exported_namespace="rhacs-kkkkllllmmmmnnnnoooo"}
        values: "0+0x379 0+4x21"
      - series: probe_success{app="rhacs", endpoint="central-ui", namespace="rhacs-aaaabbbbccccddddeeee", rhacs_instance_id="aaaabbbbccccddddeeee"}
        values: "1x1000"
      - series: probe_success{app="rhacs", endpoint="central-data", namespace="rhacs-aaaabbbbccccddddeeee", rhacs_instance_id="aaaabbbbccccddddeeee"}
        values: "1x1000"
      - series: probe_success{app="rhacs", endpoint="central-ui", namespace="rhacs-kkkkllllmmmmnnnnoooo", rhacs_instance_id="kkkkllllmmmmnnnnoooo"}
        values: "1x400"
      - series: probe_success{app="rhacs", endpoint="central-data", namespace="rhacs-kkkkllllmmmmnnnnoooo", rhacs_instance_id="kkkkllllmmmmnnnnoooo"}
        values: "1x400"
    alert_rule_test:
      - eval_time: 100m
        alertname: Central availability error budget exhaustion - 50%
        exp_alerts: []
      - eval_time: 2000m
        alertname: Central availability error budget exhaustion - 50%
        exp_alerts:
          - exp_labels:
              alertname: Central availability error budget exhaustion - 50%
              service: central
              severity: warning
              namespace: rhacs-kkkkllllmmmmnnnnoooo
              rhacs_instance_id: kkkkllllmmmmnnnnoooo
            exp_annotations:
              message: "High availability error budget exhaustion for central. Current exhaustion: 50.84%."
              sop_url: "https://gitlab.cee.redhat.com/stackrox/acs-managed-service-runbooks/blob/master/sops/dp-018-rhacs-central-slo-alerts.md"

  # Central weekly exhaustion
  - interval: 30m
    input_series:
      - series: kube_deployment_status_replicas_ready{deployment="central", namespace="rhacs-ffffgggghhhhiiiijjjj"}
        values: "0+0x5 1+0x265 0+0x15 1+0x100"
      - series: probe_success{app="rhacs", endpoint="central-ui", namespace="rhacs-ffffgggghhhhiiiijjjj", rhacs_instance_id="ffffgggghhhhiiiijjjj"}
        values: "1x400"
      - series: probe_success{app="rhacs", endpoint="central-data", namespace="rhacs-ffffgggghhhhiiiijjjj", rhacs_instance_id="ffffgggghhhhiiiijjjj"}
        values: "1x400"
    alert_rule_test:
      - eval_time: 100m
        alertname: Central availability weekly exhaustion
        exp_alerts: []
      - eval_time: 11000m
        alertname: Central availability weekly exhaustion
        exp_alerts:
          - exp_labels:
              alertname: Central availability weekly exhaustion
              service: central
              severity: warning
              namespace: rhacs-ffffgggghhhhiiiijjjj
              rhacs_instance_id: ffffgggghhhhiiiijjjj
            exp_annotations:
              message: "Availability error budget exhaustion has increased by 19.84% over the last week."
              sop_url: "https://gitlab.cee.redhat.com/stackrox/acs-managed-service-runbooks/blob/master/sops/dp-018-rhacs-central-slo-alerts.md"

  # Central high availability burn rate
  - interval: 5m
    input_series:
      # 30m downtime due to pod not ready.
      - series: kube_deployment_status_replicas_ready{deployment="central", namespace="rhacs-ppppqqqqrrrrsssstttt"}
        values: "1+0x6 0+0x6"
      - series: grpc_server_handled_total{job="central", grpc_type="unary", grpc_code="OK", namespace="rhacs-ppppqqqqrrrrsssstttt"}
        values: "5+5x14"
      - series: haproxy_backend_http_responses_total{job="router_internal_default", code="2xx", exported_namespace="rhacs-ppppqqqqrrrrsssstttt"}
        values: "5+5x14"
      - series: probe_success{app="rhacs", endpoint="central-ui", namespace="rhacs-ppppqqqqrrrrsssstttt", rhacs_instance_id="ppppqqqqrrrrsssstttt"}
        values: "1x14"
      - series: probe_success{app="rhacs", endpoint="central-data", namespace="rhacs-ppppqqqqrrrrsssstttt", rhacs_instance_id="ppppqqqqrrrrsssstttt"}
        values: "1x14"
    alert_rule_test:
      - eval_time: 30m
        alertname: Central high availability burn rate
        exp_alerts: []
      - eval_time: 69m
        alertname: Central high availability burn rate
        exp_alerts:
          - exp_labels:
              alertname: Central high availability burn rate
              service: central
              severity: critical
              namespace: rhacs-ppppqqqqrrrrsssstttt
              rhacs_instance_id: ppppqqqqrrrrsssstttt
            exp_annotations:
              message: "High availability burn rate for central. Current burn rate per hour: 57.5."
              sop_url: "https://gitlab.cee.redhat.com/stackrox/acs-managed-service-runbooks/blob/master/sops/dp-018-rhacs-central-slo-alerts.md"
