rule_files:
  - /tmp/prometheus-rules-test.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: up{namespace="rhacs", pod="fleetshard-sync-1234", instance="1.2.3.4:9090"}
        values: "0+0x20 1+0x20"
      - series: kube_pod_container_status_ready{namespace="rhacs", pod="fleetshard-sync-1234"}
        values: "1+0x40"
    alert_rule_test:
      - eval_time: 10m
        alertname: RHACSFleetshardSyncScrapeFailed
        exp_alerts: []
      - eval_time: 25m
        alertname: RHACSFleetshardSyncScrapeFailed
        exp_alerts:
          - exp_labels:
              alertname: RHACSFleetshardSyncScrapeFailed
              instance: 1.2.3.4:9090
              namespace: rhacs
              pod: fleetshard-sync-1234
              severity: critical
            exp_annotations:
              summary: "Prometheus unable to scrape metrics from target `fleetshard-sync-1234` in namespace `rhacs`."
              description: "During the last 10 minutes, only `45.45%` of scrapes of target `fleetshard-sync-1234` in namespace `rhacs` were successful. This alert is raised when less than 50% of scrapes are successful."
              sop_url: "https://gitlab.cee.redhat.com/stackrox/acs-managed-service-runbooks/blob/master/sops/dp-005-fleetshard-sync-unavailable.md"
